What is the difference between a method and a function?

<!-- grade_id: q1_ans, points: 2, grade: true, solution: true -->
```python

```

Convert the below variable to a string and assign to `string_var`. Assign the length of that string to: `len_var`

<!-- grade_id: q2_ans, points: 0, grade: false, solution: true -->
```python
initial_var = 358645317684531432678
string_var = str(initial_var)
length_var = None
print(length_var)
```

Output:
```
None
```

<!-- grade_id: q2_test, points: 3, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_not_equal

assert_equal(length_var, 21)
assert_not_equal(string_var, None)
assert_not_equal(length_var, None)
```

Output:
```
/usr/local/Cellar/python3/3.5.2/Frameworks/Python.framework/Versions/3.5/lib/python3.5/unittest/case.py in assertEqual(self, first, second, msg)
    818         """
    819         assertion_func = self._getAssertEqualityFunc(first, second)
--> 820         assertion_func(first, second, msg=msg)
    821 
    822     def assertNotEqual(self, first, second, msg=None):

/usr/local/Cellar/python3/3.5.2/Frameworks/Python.framework/Versions/3.5/lib/python3.5/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    811             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    812             msg = self._formatMessage(msg, standardMsg)
--> 813             raise self.failureException(msg)
    814 
    815     def assertEqual(self, first, second, msg=None):

AssertionError: None != 21
```

What do we mean when we say `local` scope?

<!-- grade_id: q3_ans, points: 2, grade: true, solution: true -->


What will the output of the code below be? Highlight the line numbers where the scope changes.

```python
name = "Oi, you there"

def update_name():
    name = "what?"
    
print(name)
update_name()
print(name)
```

Output:
```
Oi, you there
what?
```

<!-- grade_id: q4_ans, points: 4, grade: true, solution: true -->
```python

```

What would change if you added the line `global name` into the function `update_name`?

<!-- grade_id: q5_ans, points: 2, grade: true, solution: true -->
```python

```

Write a script below that asks a user for 2 numbers, these will not be integers by default. Pass those numbers to a function, called `power`, that returns the first argument to the power of the second. You will need to look up the symbol for `to the power of` in the Python docs

<!-- grade_id: q6_ans, points: 0, grade: false, solution: true -->
```python
num_one = None
num_two = None

def square(one, two):
    pass



```

<!-- grade_id: q6_tests, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_not_equal

assert_equal(square(12,2), 144)
assert_equal(square(50, 4), 6250000)
assert_equal(square(2,2), 4)
```